package foo;

import org.junit.Test;

/**
 * top javadoc
 */
public class SomeTest {

    /**
     * @testName: flushBufferTest
     *
     * @assertion_ids: Servlet:JAVADOC:348
     *
     * @test_Strategy: Servlet wraps response. Servlet writes data in the buffer
     * and flushes it
     */

    /*
     * @testName: getBufferSizeTest
     *
     * @assertion_ids: Servlet:JAVADOC:347
     *
     * @test_Strategy: Servlet wraps response. Servlet flushes buffer and verifies
     * the size of the buffer
     */

    /*
     * @testName: getLocaleTest
     *
     * @assertion_ids: Servlet:JAVADOC:354
     *
     * @test_Strategy: Servlet wraps response. Servlet set Locale and then
     * verifies it
     *
     */

    /*
     * @testName: testIt
     *
     * @assertion_ids: Servlet:JAVADOC:485
     *
     * @test_Strategy: just do this
     *
     */

    /**
     * here we go
     * @throws Exception
     */
    @Test
    public void testIt() throws Exception {
        // not much really
    }

    /**
     *
     *
     */
    private void foo() throws Exception {
        // yup
    }

}